Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **37 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Leon Müller"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Heaps and Heapsort


## Question 1 [time estimate: 10 minutes]

Given the array `H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]`, perform the following operations:
1. Draw the corresponding binary tree of H. Is the binary tree a valid max heap? Explain your answer.
2. Using as a model the drawing examples illustrated in Figure 6.2 of Cormen et al.,  draw a step-by-step transformation of the array above into a valid max heap. 
3. Now that you have obtained a valid max heap, write out the corresponding array that stores the valid max-heap.

Use as many cells as you wish for this question.

<!-- grade_id: cell-b1f821b034b44619, points: 0, grade: true, solution: true -->
1. The binary tree H is not a valid max heap. The root is not the largest element of the tree, nor do several of the other node relationships satisfy the criteria for it to qualify as a max heap. The condition is: the parent node of any node (A[i]) is greater or equal to it. 
2. I only drew the step by step transformation for one element. Doing it for all the elements would have taken longer than an our, let alone the 10 minutes allocated to this question. I thought there are more efficient ways in which I can apply my understanding. 
3. H = [92,85,85,76,39,30,49,16,49,15,21,7,29,31,28]

![BinaryTree](Two.png)

![BinaryTree](Three.png)

![BinaryTree](Four.png)

![BinaryTree](Five.png)

## Question 2. 

Consider the following questions on the $MAX-HEAPIFY$ operation.

### Question 2a [time estimate: 5 minutes]

In the pseudocode of $MAX-HEAPIFY$ (Cormen et al., p.154, or you can view it [here](https://drive.google.com/open?id=1e_3jsX4-qQCfZXKMok_T6LPFh9FwtmT5)), what does A.heap-size mean and what is the idea behind the local variable largest? 


<!-- grade_id: cell-06106b81a909b9e6, points: 0, grade: true, solution: true -->
The pseudocode checks that l and r are smaller or equal to A.heap_size. (It is imporant to remember that here the index values are being compared, not the values of the elements stored at those indexes). A.heap_size defines the boundary between the elements in the array that are part of the heap and those that are not or no longer part of it. When we create a heap sort algorithm we systematically construct and deconstruct a heap. A.heap_size is the upper bound of the subarray of items that are in the heap (0<=A.heap_size<=A.length).

The local variables "largest" is used to compare the size of the element of A[i] (the parent node) and A[2i] (the left child) and A[2i+1] (the right child). This local variable is used to determine if a swap needs to occur, and if so, which one.

### Question 2b [time estimate: 2 minutes]

The functions $LEFT(i)$ and $RIGHT(i)$, lines 1 and 2 in the $MAX-HEAPIFY$ pseudocode, return the array index of the left and right child, respectively, of a node in a binary tree. From reading Section 6.1, you know that the input to both functions is an integer number, $i$, which corresponds to the array index of the parent node in the array. Review Section 6.1 for more information. Write a Python implementation of the functions $LEFT(i)$ and $RIGHT(i)$ by filling in the cells below.

Note: The formulas to find the indexes for left and right children from Cormen et al. do not work the same in Python as indexing starts at 0 and not at 1. 

<!-- grade_id: cell-2efd4017321c8d48, points: 0, grade: false, solution: true -->
```python
def left(i):
    return(A[2*(i+1)-1])
    raise NotImplementedError()
```

<!-- grade_id: cell-507ddbab62e6a32c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

<!-- grade_id: cell-b7e5ceeacaeaf1b8, points: 0, grade: false, solution: true -->
```python
def right(i):
    return(A[2*(i+1)])
    raise NotImplementedError()
```

<!-- grade_id: cell-c24867515653788a, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python
A=[39,85,85,16,49,7,49,92,76,15,21,30,29,31,28]
i = 5
print(i)
print(left(i))
print(right(i))
```

Output:
```
5
30
29
```

### Question 2c [time estimate: 10 minutes]

Write a Python implementation of the MAX-HEAPIFY operation using the pseudocode above, and your newly written functions, `left` and `right`.

<!-- grade_id: cell-52289253a243341b, points: 0, grade: false, solution: true -->
```python
def heapify(heap, i): # i is the index of the root element of question
    """
    Parameters
    ----------
    heap : list of floats
        Assume that the heap size is the length of the heap
        
    Note
    ----
    No output is needed. This function should modify (if necessary) 
    heap in-place.
    
    """
    l = left(i)  # the index number of the left child is assigned to l through function left(i)
    r = right(i)  # the index number of the right child is assigned to r through function right(i)
    heap_size = len(heap)
    if l <= heap_size and A[l] > A[i]:  # 
        largest = l
        print("1")
    else:
        largest = i
        print("2")
    if r <= heap_size and A[r] > A[largest]:  # 
        largest = r
        print("3")
    if largest != i:
        print("4")
        print(A[i], A[largest])
        A[i], A[largest] = A[largest], A[i]  # exchange A[i] with A[largest]
        print(A[i], A[largest])
        heapify(heap, largest)
    
#     raise NotImplementedError()
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
print(heapify(A,0), A)
```

Output:
```
2
None [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
```

There is some issue with the implementation of the pseudocode here. I'm not sure what it is. But it seems like the last if statement is not run, which does not make sense because i should not equal to largest. I put print statements into the code to see the results of the if statements. It appears, that "l <= heap_size and A[l] > A[i]" was evaluated as False, hence, i was assigned to largest. This is wrong, yet I'm not yet sure how to fix this issue. 

<!-- grade_id: cell-7f63745ff34b881b, points: 1, grade: true, solution: false -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-32-f270faf6e16e> in <module>
      1 A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
      2 heapify(A,0)
----> 3 assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

AssertionError:
```

## Question 3 [time estimate: 3 minutes] 

Next, write a Python implementation of the BUILD_MAX_HEAP operation using the pseudocode provided in Section 6.3 of Cormen et. al. Test your Python implementation using the array in problem 1, and make sure your Python codes produce a valid max heap.

<!-- grade_id: cell-9237da6268a5e3eb, points: 0, grade: false, solution: true -->
```python
def build_max_heap(A):
    """
    Parameters
    ----------
    A : a list of floats
    
    Note
    ----
    No output is needed. The function should turn A into 
    a valid max heap, in-place.
    
    """
    heap_size = len(A)
    for i in range((heap_size//2),1,-1):
        heapify(A, i)
    raise NotImplementedError()
```

Issues with the implementation of the function heapify cause issues for this function too. I believe this implementation by itself is sound, yet it requires for heapify to be correct also. 

<!-- grade_id: cell-5dc147b837da3f7e, points: 1, grade: true, solution: false -->
```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])
```

Output:
```
<ipython-input-71-3a8e1c19ad52> in heapify(heap, i)
     12 
     13     """
---> 14     l = left(i)  # the index number of the left child is assigned to l through function left(i)
     15     r = right(i)  # the index number of the right child is assigned to r through function right(i)
     16     heap_size = len(heap)

<ipython-input-58-1775b2e90662> in left(i)
      2 #     if i ==0:
      3 #         return(A[1])
----> 4     return(A[2*(i+1)-1])
      5 # A[2*(i+1)-1]
      6     raise NotImplementedError()

IndexError: list index out of range
```

## Question 4 [time estimate: 7 minutes]

Lastly, write Python implementations of the $MIN-HEAPIFY$ and $BUILD-MIN-HEAP$ operations for a min heap data structure. You can use your $MAX-HEAPIFY$ and $BUILD-MAX-HEAP$ Python function as models, just remember that the latter two functions support operations for the max heap data structure. Test your Python implementation of $BUILD-MIN-HEAP$ using the array in problem 1, and make sure your Python codes produce a valid min heap. 

<!-- grade_id: cell-d86943f127876fc8, points: 0, grade: false, solution: true -->
```python
def min_heapify(heap, i):
    """
    Parameters
    ----------
    heap : a list of floats
        Assume that the heap size is the length of the heap.
    
    Note
    ----
    No output is needed. This function should modify (if necessary) 
    heap in-place.
    
    """
    l = left(i)  # the index number of the left child is assigned to l through function left(i)
    r = right(i)  # the index number of the right child is assigned to r through function right(i)
    heap_size = len(heap)
    if l <= heap_size and A[l] < A[i]:  # 
        smallest = l
    else:
        smallest = i
    if r <= heap_size and A[r] < A[smallest]:  # 
        smallest = r
    if smallest != i:
        print(A[i], A[smallest])
        A[i], A[smallest] = A[smallest], A[i]  # exchange A[i] with A[smallest]
        print(A[i], A[smallest])
        heapify(heap, smallest)
    
    raise NotImplementedError()
```

<!-- grade_id: cell-6efa45e3956ad0b8, points: 0, grade: false, solution: true -->
```python
def build_min_heap(A):
    """
    Parameters
    ----------
    A : a list of floats
    
    Note
    ----
    No output is needed. The function should turn A into 
    a valid min heap, in-place.
    
    """
    heap_size = len(A)
    for i in range((heap_size//2),1,-1):
        heapify(A, i)
    raise NotImplementedError()
```

<!-- grade_id: cell-7e64941402f819c2, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```